rule_files:
- ../aggregate-prometheus-rules/elasticsearch.rules.yaml

evaluation_interval: 30s

tests:
- interval: 30s
  input_series:
  # ElasticsearchDown
  - series: 'up{job="elasticsearch-logging"}'
    values: '0+0x30'
  alert_rule_test:
  - eval_time: 15m
    alertname: ElasticsearchDown
    exp_alerts:
    - exp_labels:
        service: elasticsearch-logging
        severity: warning
        type: seed
        visibility: operator
      exp_annotations:
        description: There are no running elasticsearch pods. No logs will be collected.
        summary: Elasticsearch is down
