rule_files:
- ../aggregate-prometheus-rules/fluentd.rules.yaml

evaluation_interval: 30s

tests:
- interval: 30s
  input_series:
  # FluentdDown
  - series: 'up{app="fluentd-es"}'
    values: '0+0x30'
  alert_rule_test:
  - eval_time: 15m
    alertname: FluentdDown
    exp_alerts:
    - exp_labels:
        service: fluentd-es
        severity: warning
        type: seed
        visibility: operator
      exp_annotations:
        description: There are no running fluentd pods. No logs will be collected.
        summary: Fluentd is down

